<p style="font-family: Arial; font-size:3.75vw;color:purple; font-style:bold"><br>
Regression Exercise Notebook
</p><br>

# Exercise Notebook Instructions

### 1. Important: Only modify the cells which instruct you to modify them - leave "do not modify" cells alone.  

The code which tests your responses assumes you have run the startup/read-only code exactly.

### 2. Work through the notebook in order.

Some of the steps depend on previous, so you'll want to move through the notebook in order.

### 3. It is okay to use numpy libraries.

You may find some of these questions are fairly straightforward to answer using built-in numpy functions.  That's totally okay - part of the point of these exercises is to familiarize you with the commonly used numpy functions.

### 4. Seek help if stuck

If you get stuck, don't worry!  You can either review the videos/notebooks from this week, ask in the course forums, or look to the solutions for the correct answer.  BUT, be careful about looking to the solutions too quickly.  Struggling to get the right answer is an important part of the learning process.

```python
# DO NOT MODIFY

# import appropriate libraries
import sqlite3
import pandas as pd 
from sklearn.tree import DecisionTreeRegressor
from sklearn.linear_model import LinearRegression
from sklearn.model_selection import train_test_split
from sklearn.metrics import mean_squared_error
from math import sqrt
import numpy as np
from sklearn.model_selection import train_test_split
```

```python
# DO NOT MODIFY

# We will use the European Soccer dataset for this exercise.

def get_data():
    cnx = sqlite3.connect('./database.sqlite')
    df = pd.read_sql_query("SELECT * FROM Player_Attributes", cnx)
    return df

df = get_data()
```

<p style="font-family: Arial; font-size:2.75vw;color:purple; font-style:bold"><br>

Exercise 1: Drop NULLs in the Data<br><br></p>


In the cell below, modify the function `preparation`. The `preparation` function takes three arguments:
* a dataframe, 
* a list of features, and 
* the name of regression target column.

Function should do the following: 
- to take the input data frame and remove all rows containing NULLs. 
- RETURN two data frames, one containing the feature columns and other the target column

<!-- grade_id: c555, points: 0, grade: false, solution: true -->
```python
# modify this cell

def preparation(dataf, features, target):
    ### BEGIN SOLUTION
    dataf.dropna(inplace=True)
    featuredf = dataf[features]
    targetdf = dataf[target]
    return featuredf,targetdf
    ### END SOLUTION
```

```python
features = ['potential', 'reactions', 'vision']
target = ['overall_rating']
X, y = preparation(df, features, target)
X.shape, y.shape
```

Output:
```
((180354, 3), (180354, 1))
```

<!-- grade_id: ex15, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY
ans1 = ['potential', 'reactions', 'vision']
ans2 = ['overall_rating']

try: 
    features = ans1
    target   = ans2
    X, y = preparation(df, features, target)
    X.columns, y.columns
    assert np.alltrue(X.columns.tolist() == ans1)
    assert np.alltrue(y.columns.tolist() == ans2)
    assert np.alltrue(X.shape[0] == 180354)
    assert np.alltrue(y.shape[0] == 180354)

except AssertionError as e: print("Try again, your output did not match the expected answer above")
```

<p style="font-family: Arial; font-size:2.75vw;color:purple; font-style:bold"><br>

Exercise 2: Perform Splitting<br><br></p>


In the cell below, modify the function to take features (X) and target (y) dataframe and 
split 70% as training data and 30% as test data, using a random state = rstate.

The function should return X_train, X_test, y_train, and y_test

<!-- grade_id: 5551, points: 0, grade: false, solution: true -->
```python
# modify this cell

def clean_data(X, y, rstate):
    ### BEGIN SOLUTION
    X_train, X_test, y_train, y_test = train_test_split(X,y,test_size=0.33,random_state=rstate)
    return X_train, X_test, y_train, y_test
    ### END SOLUTION
```

```python
X_train, X_test, y_train, y_test = clean_data(X, y, 9000)
print(X_train.shape, y_train.shape)
print(X_test.shape, y_test.shape)
```

Output:
```
C:\Users\Public\Anaconda3\lib\site-packages\sklearn\utils\validation.py in <listcomp>(.0)
    199     """
    200 
--> 201     lengths = [_num_samples(X) for X in arrays if X is not None]
    202     uniques = np.unique(lengths)
    203     if len(uniques) > 1:

C:\Users\Public\Anaconda3\lib\site-packages\sklearn\utils\validation.py in _num_samples(x)
    144         if len(x.shape) == 0:
    145             raise TypeError("Singleton array %r cannot be considered"
--> 146                             " a valid collection." % x)
    147         # Check that shape is returning an integer or default to len
    148         # Dask dataframes may not return numeric shape[0] value

TypeError: Singleton array array(9000) cannot be considered a valid collection.
```

<!-- grade_id: ex151, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY

try: 
    X_train, X_test, y_train, y_test = clean_data(X, y, 9000)

    assert np.alltrue(X_train.shape == (126247, 3))
    assert np.alltrue(y_train.shape == (126247, 1))
    assert np.alltrue(X_test.shape == (54107, 3))
    assert np.alltrue(y_test.shape == (54107, 1))

except AssertionError as e: print("Try again")
```

<p style="font-family: Arial; font-size:2.75vw;color:purple; font-style:bold"><br>

Exercise 3: Build a Regressor<br><br></p>

In the cell below, modify the function to take X_train, y_train only and RETURN a regressor
for predicting the y_train based on columns in X_train. You can pick any regressor model.

The function should RETURN a trained model. We will test your regressor on X_test and y_test

<!-- grade_id: c5553, points: 0, grade: false, solution: true -->
```python
# modify this cell

def train_regressor(X_train, y_train):
    ### BEGIN SOLUTION

    ### END SOLUTION
```

<!-- grade_id: ex153, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY

threshold = 4.5

try: 
    model = train_regressor(X_train, y_train['overall_rating'])
    y_prediction = model.predict(X_test)
    rmse = sqrt(mean_squared_error(y_true = y_test, y_pred = y_prediction))
    #print(rmse)
    assert np.alltrue(rmse < threshold)
except AssertionError as e: print("Keep trying - can you get an RMSE < %f" % threshold)
```